# Lab-1 - Some Basic Growth Facts

## Purpose

This introductory lab will serve two purposes.  

1. It will acquaint you with some basic growth facts that you will seek to explain later in the semester.  
2. It will allow you to apply your new knowledge of Python and Jupyter notebooks.

# Assignment

In this lab, you will document cross-country differences in income and growth rates to help motivate the study of growth throughout the semester. The final product is a Jupyter notebook that provides and answer to all the questions. It must perform all analyses and create all the figures, and export the outputs requested. 

# Setup -  Import packages and setup paths

Create a `pathout` variable and `pathgraphs` for the name of the directories where you will save data and figures. If they do not exist, make sure to create them

<!-- grade_id: cell-1f18136b34393c80, points: 0, grade: false, solution: true -->
```python
import os
import matplotlib.pyplot as plt
import numpy as np
import pandas as pd

pathout = './data/'
pathgraphs = "./graphs/"
os.makedirs(pathout, exist_ok = True)
os.makedirs(pathgraphs, exist_ok = True)
```

<!-- grade_id: cell-b1aa21e6c45d7cea, points: 2.5, grade: true, solution: false -->
```python
assert(os.path.exists(pathout)==True)
assert(os.path.exists(pathgraphs)==True)
```

# 1. Import Data

Use the code from class to create income per capita data from the Penn World Tables and the World Development Indicators for the 1950-2021 period.

Create a dataframe for each data source that contains the following variables:
1. Country code: `country_code`
2. Country name: `country_name`
2. Year: `year`
3. GDP per capita: `gdp_pc`

For PWT use only `rgdpe` for GDP and name the dataframe `pwt`, for WDi use `'NY.GDP.PCAP.PP.KD'` for GDP per capita and name it `wdi`.

Save the files with names `pwt100.dta` and `wdi.dta` in the `pathout` directory.

## PWT-10

<!-- grade_id: cell-106b8ca8b19737cb, points: 0, grade: false, solution: true -->
```python
data = pd.read_excel("/Users/colinduggan/Desktop/firstsemester/ECO4362/labs/lab-1-ColinDuggan02/data/PWT1.xlsx", 'Data')
regions = pd.read_excel("/Users/colinduggan/Desktop/firstsemester/ECO4362/labs/lab-1-ColinDuggan02/data/PWT1.xlsx", 'Regions')
#data.head()
#regions.head()
regions.drop(regions.columns[[0]], axis=1)

PWT = data.merge(regions, left_on='RegionCode', right_on='IsoAlpha3')
PWT = PWT.drop('Unnamed: 0', axis = 1)
PWT = PWT[PWT['VariableCode'] == 'rgdpe']
PWT = PWT.rename({'RegionCode':'country_code', 'Name':'country_name', 'YearCode':'year', 'AggValue':'gdp_pc'}, axis=1)


PWT.head()
# Download PWT-10 data
# YOUR CODE HERE

# Generate GDPpc using rgdpe measure
# YOUR CODE HERE


# Select columns and rename them if necessary
# YOUR CODE HERE


# Save dataframe with 4 requested variables
# YOUR CODE HERE

```

Output:
```
VariableCode country_code  year       gdp_pc IsoAlpha3 Code country_name
0        rgdpe          ABW  1970   400.081726       ABW   AW        Aruba
2        rgdpe          ABW  1980   955.632141       ABW   AW        Aruba
4        rgdpe          ABW  1990  2575.255615       ABW   AW        Aruba
6        rgdpe          ABW  2000  4031.134277       ABW   AW        Aruba
8        rgdpe          ABW  2010  4012.711670       ABW   AW        Aruba
```

<!-- grade_id: cell-9be909839ff1dbde, points: 5, grade: true, solution: false -->
```python
assert(pwt.shape==(12810,4))
assert(os.path.exists(pathout + '/pwt100.dta')==True)
```

## WDI

Get countries and their regions and basic info useful later, especially for plotting

<!-- grade_id: cell-32c051ffcc13b92d, points: 0, grade: false, solution: true -->
```python
wdi = pd.read_excel("/Users/colinduggan/Desktop/firstsemester/ECO4362/labs/lab-1-ColinDuggan02/data/GDPpc.xls", "Data")
wdi = wdi.drop([0,1])
wdi.columns = wdi.iloc[0]
wdi = wdi.drop(2)

for i in range(1960, 2022):
    if np.isnan(wdi[float(i)].mean()):
        wdi = wdi.drop(float(i), axis = 1) #dropping rows with mean as NaN

wdi

```

Output:
```
2                   Country Name Country Code  \
3                          Aruba          ABW   
4    Africa Eastern and Southern          AFE   
5                    Afghanistan          AFG   
6     Africa Western and Central          AFW   
7                         Angola          AGO   
..                           ...          ...   
264                       Kosovo          XKX   
265                  Yemen, Rep.          YEM   
266                 South Africa          ZAF   
267                       Zambia          ZMB   
268                     Zimbabwe          ZWE   

2                                       Indicator Name     Indicator Code  \
3    GDP per capita, PPP (constant 2017 internation...  NY.GDP.PCAP.PP.KD   
4    GDP per capita, PPP (constant 2017 internation...  NY.GDP.PCAP.PP.KD   
5    GDP per capita, PPP (constant 2017 internation...  NY.GDP.PCAP.PP.KD   
6    GDP per capita, PPP (constant 2017 internation...  NY.GDP.PCAP.PP.KD   
7    GDP per capita, PPP (
```

Download data

<!-- grade_id: cell-a8f3321fbe617376, points: 0, grade: false, solution: true -->
```python

```

Merge WDI datasets so we have codes and names together

<!-- grade_id: cell-87ee6c45fa2f04e5, points: 0, grade: false, solution: true -->
```python

```

Select the 4 requested variables and save file

<!-- grade_id: cell-5993e1eaa1dea3d6, points: 0, grade: false, solution: true -->
```python

```

<!-- grade_id: cell-ed4011b449d77532, points: 5, grade: true, solution: false -->
```python
assert(wdi.shape==(13454,4))
assert(os.path.exists(pathout + '/wdi.dta')==True)
assert((wdi.columns==['country_code', 'country_name', 'year', 'gdp_pc']).all())
```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [157], in <cell line: 1>()
----> 1 assert(wdi.shape==(13454,4))
      2 assert(os.path.exists(pathout + '/wdi.dta')==True)
      3 assert((wdi.columns==['country_code', 'country_name', 'year', 'gdp_pc']).all())

AssertionError:
```

# 2. Compare Income levels

For each of these two datasets: 

* compute the average level of real per capita income in the earliest and in the latest year across all the countries in each dataset.  
* Find the minimum per capita income in each of these two periods
* Compute the maximum per capita income in each of these two periods
* Compute the ratio of GDP per capita between richest to poorest in the two periods
* Show how the distribution of income has changed over time by  plotting the distributions of log-income per capita in different years using kernel densities. Export each figure as a pdf file into ``pathgraphs`` with name ``dataset_gdppc_year.pdf``. To do this, write a function that takes as inputs the dataset and the year you want to plot.
* Create a figure that compares the distribution of log-income per capita in both years for each dataset, i.e., it shows both jointly in one figure, and export it as a pdf and a png file into ``pathgraphs`` with names ``dataset_gdppc_minyear_maxyear.pdf`` and ``dataset_gdppc_minyear_maxyear.png``.

You need to name the variables in a way that the example code runs. 

Save the earliest year in PWT as `year_min_pwt`.

<!-- grade_id: cell-afc2623d777d61a3, points: 0, grade: false, solution: true -->
```python
year_min_pwt = PWT['year'].min()
```

<!-- grade_id: cell-f5ddd6c5017b35c7, points: 1, grade: true, solution: false -->
```python
print("Earliest year in PWT is ", year_min_pwt, '\n')
```

Output:
```
Earliest year in PWT is  1950
```

Earliest year in WDI as `year_min_wdi`

<!-- grade_id: cell-4b1a8c46621f3029, points: 0, grade: false, solution: true -->
```python
year_min_wdi = 1990.0
```

<!-- grade_id: cell-e567b309b46cab0d, points: 1, grade: true, solution: false -->
```python
print("Earliest year in WDI is ", year_min_wdi, '\n')
```

Output:
```
Earliest year in WDI is  1990.0
```

Latest year in PWT as `year_max_pwt`

<!-- grade_id: cell-13da7bcc2b029d4b, points: 0, grade: false, solution: true -->
```python
year_max_pwt = PWT['year'].max()
```

<!-- grade_id: cell-c87de0e2fbfc491c, points: 1, grade: true, solution: false -->
```python
print("Latest year in PWT is ", year_max_pwt, '\n')
```

Output:
```
Latest year in PWT is  2017
```

Latest year in WDI as `year_max_wdi`

<!-- grade_id: cell-0f3090c4ac03a294, points: 0, grade: false, solution: true -->
```python
year_max_wdi = 2021.0
```

<!-- grade_id: cell-ffbf9548f485bccd, points: 1, grade: true, solution: false -->
```python
print("Latest year in WDI is ", year_max_wdi, '\n')
```

Output:
```
Latest year in WDI is  2021.0
```

Average GDP pc in earliest year in PWT as `mean_gdppc_pwt_early`

<!-- grade_id: cell-5e01577356a7efd9, points: 0, grade: false, solution: true -->
```python
mean_gdppc_pwt_early = PWT[PWT['year'] == year_min_pwt].mean()
```

Output:
```
/var/folders/ft/0z4sd94x4lvgyjfx2r74bmf00000gn/T/ipykernel_1696/3931691114.py:1: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  mean_gdppc_pwt_early = PWT[PWT['year'] == year_min_pwt].mean()
```

<!-- grade_id: cell-2e11fc7a01cc7420, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", mean_gdppc_pwt_early, '\n')
```

Output:
```
The average income per capita in the earliest year, i.e., 1950 in the PWT, is  year        1950.000000
gdp_pc    112895.837356
dtype: float64
```

Average GDP pc in earliest year in WDI as `mean_gdppc_wdi_early`

<!-- grade_id: cell-a008aac4b29d6ed8, points: 0, grade: false, solution: true -->
```python
mean_gdppc_wdi_early = wdi[year_min_wdi].mean()
```

<!-- grade_id: cell-62297975f3f5bf56, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", mean_gdppc_wdi_early, '\n')
```

Output:
```
The average income per capita in the earliest year, i.e., 1990.0 in the WDI, is  12465.9416279256
```

Average GDP pc in latest year in PWT as `mean_gdppc_pwt_late`

<!-- grade_id: cell-793af247164dac85, points: 0, grade: false, solution: true -->
```python
mean_gdppc_pwt_late = PWT[PWT['year'] == year_max_pwt].mean()
```

Output:
```
/var/folders/ft/0z4sd94x4lvgyjfx2r74bmf00000gn/T/ipykernel_1696/137744828.py:1: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  mean_gdppc_pwt_late = PWT[PWT['year'] == year_max_pwt].mean()
```

<!-- grade_id: cell-c271b86663463a1e, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", mean_gdppc_pwt_late, '\n')
```

Output:
```
The average income per capita in the latest year, i.e., 2017 in the PWT, is  year        2017.000000
gdp_pc    652598.971887
dtype: float64
```

Average GDP pc in latest year in WDI as `mean_gdppc_wdi_late`

<!-- grade_id: cell-865d23fd2e756892, points: 0, grade: false, solution: true -->
```python
mean_gdppc_wdi_late = wdi[year_max_wdi].mean()
```

<!-- grade_id: cell-23fce009562d37a6, points: 1, grade: true, solution: false -->
```python
print("The average income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", mean_gdppc_wdi_late, '\n')
```

Output:
```
The average income per capita in the latest year, i.e., 2021.0 in the WDI, is  19925.125287380168
```

Minimum GDP pc in earliest year in PWT as `min_gdppc_pwt_early`

<!-- grade_id: cell-0e4cd8962831097b, points: 0, grade: false, solution: true -->
```python
min_gdppc_pwt_early = PWT[PWT['year']==year_min_pwt].min()['gdp_pc']
```

<!-- grade_id: cell-c203c74cc486219b, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", min_gdppc_pwt_early, '\n')
```

Output:
```
The minimum income per capita in the earliest year, i.e., 1950 in the PWT, is  1293.68737792969
```

Minimum GDP pc in earliest year in WDI as `min_gdppc_wdi_early`

<!-- grade_id: cell-c1ececdca468482c, points: 0, grade: false, solution: true -->
```python
min_gdppc_wdi_early = wdi[year_min_wdi].min()
```

<!-- grade_id: cell-818e66b92095b995, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", min_gdppc_wdi_early, '\n')
```

Output:
```
The minimum income per capita in the earliest year, i.e., 1990.0 in the WDI, is  471.3251025477527
```

Minimum GDP pc in latest year in PWT as `min_gdppc_pwt_late`

<!-- grade_id: cell-ef74288b9b6dbe8c, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
min_gdppc_pwt_late = PWT[PWT['year']==year_max_pwt].min()['gdp_pc']
```

Output:
```
/var/folders/ft/0z4sd94x4lvgyjfx2r74bmf00000gn/T/ipykernel_1696/1791470492.py:2: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  min_gdppc_pwt_late = PWT[PWT['year']==year_max_pwt].min()['gdp_pc']
```

<!-- grade_id: cell-0cf9e07785a0b5a5, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", min_gdppc_pwt_late, '\n')
```

Output:
```
The minimum income per capita in the latest year, i.e., 2017 in the PWT, is  90.2846908569336
```

Minimum GDP pc in latest year in WDI as `min_gdppc_wdi_late`

<!-- grade_id: cell-41938ef77dde4117, points: 0, grade: false, solution: true -->
```python
min_gdppc_wdi_late = wdi[year_max_wdi].min()
```

<!-- grade_id: cell-b7fd8dc0bab0a6bb, points: 1, grade: true, solution: false -->
```python
print("The minimum income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", min_gdppc_wdi_late, '\n')
```

Output:
```
The minimum income per capita in the latest year, i.e., 2021.0 in the WDI, is  722.0462734906972
```

Maximum GDP pc in earliest year in PWT as `max_gdppc_pwt_early`

<!-- grade_id: cell-87fbacd7d7fad503, points: 0, grade: false, solution: true -->
```python
max_gdppc_pwt_early = PWT[PWT['year']==year_min_pwt].max()['gdp_pc']
```

<!-- grade_id: cell-d0d9f60799628d67, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", max_gdppc_pwt_early, '\n')
```

Output:
```
The maximum income per capita in the earliest year, i.e., 1950 in the PWT, is  2447945.5
```

Maximum GDP pc in earliest year in WDI as `max_gdppc_wdi_early`

<!-- grade_id: cell-0c5a8189069a91bf, points: 0, grade: false, solution: true -->
```python
max_gdppc_wdi_early = wdi[year_min_wdi].max()
```

<!-- grade_id: cell-ade2c564e718253e, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", max_gdppc_wdi_early, '\n')
```

Output:
```
The maximum income per capita in the earliest year, i.e., 1990.0 in the WDI, is  111454.05566540676
```

Maximum GDP pc in latest year in PWT as `max_gdppc_pwt_late`

<!-- grade_id: cell-d3988950ccb52994, points: 0, grade: false, solution: true -->
```python
max_gdppc_pwt_late = PWT[PWT['year']==year_max_pwt].max()['gdp_pc']
```

Output:
```
/var/folders/ft/0z4sd94x4lvgyjfx2r74bmf00000gn/T/ipykernel_1696/29708583.py:1: FutureWarning: Dropping of nuisance columns in DataFrame reductions (with 'numeric_only=None') is deprecated; in a future version this will raise TypeError.  Select only valid columns before calling the reduction.
  max_gdppc_pwt_late = PWT[PWT['year']==year_max_pwt].max()['gdp_pc']
```

<!-- grade_id: cell-33296a0e0170e321, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", max_gdppc_pwt_late, '\n')
```

Output:
```
The maximum income per capita in the latest year, i.e., 2017 in the PWT, is  19754754.0
```

Maximum GDP pc in latest year in WDI as `max_gdppc_wdi_late`

<!-- grade_id: cell-73fa3e093dd67a94, points: 0, grade: false, solution: true -->
```python
max_gdppc_wdi_late = wdi[year_max_wdi].max()
```

<!-- grade_id: cell-2246ac42b3c5d740, points: 1, grade: true, solution: false -->
```python
print("The maximum income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", max_gdppc_wdi_late, '\n')
```

Output:
```
The maximum income per capita in the latest year, i.e., 2021.0 in the WDI, is  118680.04372882558
```

Ratio of GDP pc in earliest year in PWT as `ratio_gdppc_pwt_early`

<!-- grade_id: cell-05c62cd437c0eb72, points: 0, grade: false, solution: true -->
```python
ratio_gdppc_pwt_early = max_gdppc_pwt_early/min_gdppc_pwt_early
```

<!-- grade_id: cell-fa9d76c13546ff92, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the earliest year, i.e.,", year_min_pwt, "in the PWT, is ", ratio_gdppc_pwt_early, '\n')
```

Output:
```
The ratio income per capita in the earliest year, i.e., 1950 in the PWT, is  1892.223377735577
```

Ratio of GDP pc in earliest year in WDI as `ratio_gdppc_wdi_early`

<!-- grade_id: cell-bccf344ade5fdad4, points: 0, grade: false, solution: true -->
```python
ratio_gdppc_wdi_early = max_gdppc_wdi_early/min_gdppc_wdi_early
```

<!-- grade_id: cell-c3d2fc191d17efeb, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the earliest year, i.e.,", year_min_wdi, "in the WDI, is ", ratio_gdppc_wdi_early, '\n')
```

Output:
```
The ratio income per capita in the earliest year, i.e., 1990.0 in the WDI, is  236.46959405077453
```

Ratio of GDP pc in latest year in PWT as `ratio_gdppc_pwt_late`

<!-- grade_id: cell-ba9fa189cc5558c7, points: 0, grade: false, solution: true -->
```python
ratio_gdppc_pwt_late = max_gdppc_pwt_late/min_gdppc_pwt_late
```

<!-- grade_id: cell-4d114926f1ca9a5c, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the latest year, i.e.,", year_max_pwt, "in the PWT, is ", ratio_gdppc_pwt_late, '\n')
```

Output:
```
The ratio income per capita in the latest year, i.e., 2017 in the PWT, is  218805.13531694608
```

Ratio of GDP pc in latest year in WDI as `ratio_gdppc_wdi_late`

<!-- grade_id: cell-0ad7da478f2c215c, points: 0, grade: false, solution: true -->
```python
ratio_gdppc_wdi_late = max_gdppc_wdi_late/min_gdppc_wdi_late
```

<!-- grade_id: cell-4c4e2928a6ee7bb5, points: 1, grade: true, solution: false -->
```python
print("The ratio income per capita in the latest year, i.e.,", year_max_wdi, "in the WDI, is ", ratio_gdppc_wdi_late, '\n')
```

Output:
```
The ratio income per capita in the latest year, i.e., 2021.0 in the WDI, is  164.3662575184728
```

<!-- grade_id: cell-5a54b444406bec1b, points: 5, grade: true, solution: false -->
```python
assert(year_min_pwt==1950)
assert(year_max_pwt==2019)

```

Output:
```
AssertionError                            Traceback (most recent call last)
Input In [334], in <cell line: 2>()
      1 assert(year_min_pwt==1950)
----> 2 assert(year_max_pwt==2019)

AssertionError:
```

Let's now create the function to plot each distribution in each year and dataset

<!-- grade_id: cell-aa7cb53e6afe319f, points: 0, grade: false, solution: true -->
```python
sns.set(rc={'figure.figsize':(11.7,8.27)})
#sns.reset_orig()
sns.set_context("talk")

def my_kde_plots(dataset='pwt', myyear=year_min_pwt):
    '''
    This function plots the KDE-plot of log income per capita, where
    
    dataset: is either pwt or wdi (one of the two datasets you created above)
    myyear: is one of the year variables you created above, e.g., year_min_pwt, year_min_wdi
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-6e064c037304d5a8, points: 5, grade: true, solution: false -->
```python
# KDE plot for PWT in earliest period
my_kde_plots()
```

Output:
```
NameError                                 Traceback (most recent call last)
Input In [371], in <cell line: 2>()
      1 # KDE plot for PWT in earliest period
----> 2 my_kde_plots()

NameError: name 'my_kde_plots' is not defined
```

<!-- grade_id: cell-daf7e60815921e4b, points: 5, grade: true, solution: false -->
```python
# KDE plot for PWT in latest period
my_kde_plots(myyear=year_max_pwt)
```

<!-- grade_id: cell-2244a0c19379d0a7, points: 5, grade: true, solution: false -->
```python
# KDE plot for WDI in earliest period
my_kde_plots(dataset='wdi', myyear=year_min_wdi)
```

<!-- grade_id: cell-bf4af7a308022573, points: 5, grade: true, solution: false -->
```python
# KDE plot for PWT in latest period
my_kde_plots(dataset='wdi', myyear=year_max_wdi)
```

<!-- grade_id: cell-9e89765762b26403, points: 2.5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + '/pwt_gdppc_' + str(year_min_pwt) + '.pdf')==True)
assert(os.path.exists(pathgraphs + '/wdi_gdppc_' + str(year_min_wdi) + '.pdf')==True)
```

Let's write a function that plots for each dataset and pair of years the distributions in that year and exports it to both png and pdf.

<!-- grade_id: cell-20032e775009fa97, points: 0, grade: false, solution: true -->
```python
sns.set(rc={'figure.figsize':(11.7,8.27)})
#sns.reset_orig()
sns.set_context("talk")

def my_kde_plots_joint(dataset='pwt', myyears=[year_min_pwt, year_max_pwt]):
    '''
    This function plots the KDE-plot of log income per capita in two years, where
    
    dataset: is either pwt or wdi (one of the two datasets you created above)
    myyears: is a list of years
    '''
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-f93b853f045b331d, points: 2.5, grade: true, solution: false -->
```python
# Plot PWT
my_kde_plots_joint()
```

<!-- grade_id: cell-d7ceffb95f6560bf, points: 2.5, grade: true, solution: false -->
```python
# Plot WDI
my_kde_plots_joint(dataset='wdi', myyears=[year_min_wdi, year_max_wdi])
```

<!-- grade_id: cell-ef075b38ad1899b7, points: 2.5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'pwt_gdppc_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf')==True)
assert(os.path.exists(pathgraphs + 'wdi_gdppc_' + str(year_min_wdi) + '_' + str(year_max_wdi) + '.pdf')==True)
```

# 3. Growth rates

Compute the growth rate of real per capita income over the time period between the first year and the last year as

$$
growth = \log(gdp\_pc_{last}) - \log(gdp\_pc_{first})
$$

Save the results in a new dataframe called `growth_dataset`.  What is the average growth rate for all the countries in your data? Save it as `mean_growth_dataset`. What is the minimum growth rate? Save it as `min_growth_dataset`. What is the maximum growth rate? Save it as `max_growth_dataset`.

Compute growth rates

<!-- grade_id: cell-d1c97f59295e8c0b, points: 0, grade: false, solution: true -->
```python
growth_rate_pwt = np.log(PWT[PWT['year'] == year_max_pwt]['gdp_pc'])-np.log(PWT[PWT['year'] == year_min_pwt]['gdp_pc'])
#number of rows between max_pwt and min_pwt are different, yields nan when subtracting logs between the two columns
growth_rate_wdi = np.log(wdi[year_max_wdi]) - np.log(wdi[year_min_wdi])

```

Compute Average growth rate

<!-- grade_id: cell-549f20d20cf21923, points: 0, grade: false, solution: true -->
```python
mean_growth_pwt = growth_rate_pwt.mean()
mean_growth_wdi = growth_rate_wdi.mean()
```

<!-- grade_id: cell-f14f47186f4c69f6, points: 2.5, grade: true, solution: false -->
```python
print('Average growth rate between', year_min_pwt, 'and', year_max_pwt, 'according to PWT was {0:.2%}'.format(mean_growth_pwt))
print('Average growth rate between', year_min_wdi, 'and', year_max_wdi, 'according to WDI was {0:.2%}'.format(mean_growth_wdi))

```

Output:
```
Average growth rate between 1950 and 2017 according to PWT was nan%
Average growth rate between 1990.0 and 2021.0 according to WDI was 52.86%
```

Compute minimum growth rate

<!-- grade_id: cell-d56ff2ee06d45c3e, points: 0, grade: false, solution: true -->
```python
 min_growth_wdi = growth_rate_wdi.min()
```

<!-- grade_id: cell-55f6b4e018482c00, points: 0, grade: false, solution: true -->
```python
min_growth_pwt = growth_rate_pwt.min()
```

Compute maximum growth rate

<!-- grade_id: cell-d5b1c547fe35cb21, points: 0, grade: false, solution: true -->
```python
max_growth_pwt = growth_rate_pwt.max()
max_growth_wdi = growth_rate_wdi.max()
```

<!-- grade_id: cell-0f079ace95a009a2, points: 2.5, grade: true, solution: false -->
```python
print('Maximum growth rate between', year_min_pwt, 'and', year_max_pwt, 'according to PWT was {0:.2%}'.format(max_growth_pwt))
print('Maximum growth rate between', year_min_wdi, 'and', year_max_wdi, 'according to WDI was {0:.2%}'.format(max_growth_wdi))

```

Output:
```
Maximum growth rate between 1950 and 2017 according to PWT was nan%
Maximum growth rate between 1990.0 and 2021.0 according to WDI was 276.42%
```

# 4. Growth rates by WB region

Consider the regions defined by the World Bank: East Asia & Pacific, Europe & Central Asia, Latin America & Caribbean, Middle East & North Africa, North America, South Asia, and Sub-Saharan Africa.  Compute the minimum, maximum, and mean growth rates of income per capita for countries in each of these groups in the earliest and latest year in each dataset.  What patterns can you identify in the growth of income and in the dispersion of income in these regions?

Start by merging the `wbcountries` dataframe with the others to get the regional variable. Create new dataframes `growth_dataset_regions`. 

**Note:** There are 2 countries in PWT that do not appear in the `wbcountries` dataframe. 
They both belong to Latin America & Caribbean (codes: `AIA`, `MSR`). So, they will be dropped in the `joint` merge.

<!-- grade_id: cell-633e45e6582f63b8, points: 0, grade: false, solution: true -->
```python
EastAsiaPacific = ("Aruba", "Africa")

#stuck

if wdi.column("Country Name") EastAsiaPacific:
    print("1")
#wbcountries
#wdi.merge(wbcountries,left_on="Country Code", right_on="")
```

Output:
```
Input In [383]
    if wdi.column("Country Name") is in EastAsiaPacific:
                                     ^
SyntaxError: invalid syntax
```

<!-- grade_id: cell-21ad096074fb87b3, points: 2.5, grade: true, solution: false -->
```python
assert(growth_wdi.shape[0]==growth_wdi_regions.shape[0])
assert((growth_pwt.shape[0] - 2)==growth_pwt_regions.shape[0])
```

Compute the statistics by region and save them in a new dataframe `statistic_growth_dataset_regions` that has 1 column with the regions and another with the growth rates.

<!-- grade_id: cell-0bd781c17aa480fa, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

# 5. Counter-factual Analysis

Using the PWT, take the maximum growth rate that you computed in question 3 and apply it to the poorest country in the earliest date.  Then take the minimum growth rate that you computed in question 3 and apply it to the richest country in the earliest date.  How would these different growth rates affect the level of income in these countries in the last period?  Explain how this computation demonstrates the importance of the study of long-run growth.

Compute counter factual income per capita for the poorest country. Safe value in `poor_counter_factual`.

<!-- grade_id: cell-e4034c55f1e16621, points: 0, grade: false, solution: true -->
```python
poor_counter_factual =  PWT[PWT[max_growth_wdi]==year_max_pwt].max()['gdp_pc']
print(poor_counter_factual)
```

Output:
```
File ~/opt/anaconda3/envs/EconGrowthUG/lib/python3.9/site-packages/pandas/core/frame.py:3505, in DataFrame.__getitem__(self, key)
   3503 if self.columns.nlevels > 1:
   3504     return self._getitem_multilevel(key)
-> 3505 indexer = self.columns.get_loc(key)
   3506 if is_integer(indexer):
   3507     indexer = [indexer]

File ~/opt/anaconda3/envs/EconGrowthUG/lib/python3.9/site-packages/pandas/core/indexes/base.py:3623, in Index.get_loc(self, key, method, tolerance)
   3621     return self._engine.get_loc(casted_key)
   3622 except KeyError as err:
-> 3623     raise KeyError(key) from err
   3624 except TypeError:
   3625     # If we have a listlike key, _check_indexing_error will raise
   3626     #  InvalidIndexError. Otherwise we fall through and re-raise
   3627     #  the TypeError.
   3628     self._check_indexing_error(key)

KeyError: 2.764174388138774
```

Compute counter factual income per capita for the richest country. Safe value in `rich_counter_factual`.

<!-- grade_id: cell-937c935aa8cba799, points: 0, grade: false, solution: true -->
```python
rich_counter_factual = PWT[PWT[min_growth_wdi]==year_min_pwt].min()['gdp_pc']
print(rich_counter_factual)
```

Output:
```
File ~/opt/anaconda3/envs/EconGrowthUG/lib/python3.9/site-packages/pandas/core/frame.py:3505, in DataFrame.__getitem__(self, key)
   3503 if self.columns.nlevels > 1:
   3504     return self._getitem_multilevel(key)
-> 3505 indexer = self.columns.get_loc(key)
   3506 if is_integer(indexer):
   3507     indexer = [indexer]

File ~/opt/anaconda3/envs/EconGrowthUG/lib/python3.9/site-packages/pandas/core/indexes/base.py:3623, in Index.get_loc(self, key, method, tolerance)
   3621     return self._engine.get_loc(casted_key)
   3622 except KeyError as err:
-> 3623     raise KeyError(key) from err
   3624 except TypeError:
   3625     # If we have a listlike key, _check_indexing_error will raise
   3626     #  InvalidIndexError. Otherwise we fall through and re-raise
   3627     #  the TypeError.
   3628     self._check_indexing_error(key)

KeyError: -0.4952755108615179
```

<!-- grade_id: cell-55559258d246fec5, points: 5, grade: true, solution: false -->
```python
assert(rich_counter_factual<poor_counter_factual)
```

Output:
```
NameError                                 Traceback (most recent call last)
Input In [390], in <cell line: 1>()
----> 1 assert(rich_counter_factual<poor_counter_factual)

NameError: name 'rich_counter_factual' is not defined
```

# 6. Who's who in economic growth?

Sort the data by the growth rates you computed in number 3.  Identify the 10 fastest growing and the 10 slowest growing countries.  Do the fast growers have anything in common?  Do the slow growers?  Do these lists suggest potential ideas for investigation of the determinants of growth?

Let's sort and show the top 10 slowest growing countries in PWT

<!-- grade_id: cell-b98c4ebf00f67a32, points: 0, grade: false, solution: true -->
```python
growth_rate_pwt.sort_values()
growth_rate_pwt.head(10)
```

Output:
```
10   NaN
22   NaN
34   NaN
46   NaN
58   NaN
60   NaN
74   NaN
82   NaN
94   NaN
96   NaN
Name: gdp_pc, dtype: float64
```

Now in WDI

<!-- grade_id: cell-6fd548697e9f0ec1, points: 0, grade: false, solution: true -->
```python
growth_rate_wdi.sort_values()
growth_rate_wdi.head(10)
```

Output:
```
3          NaN
4     0.142516
5          NaN
6     0.371162
7     0.039406
8     1.101302
9          NaN
10    0.348977
11         NaN
12    0.419030
dtype: float64
```

Let's sort and show the top 10 fastest growing countries in PWT

<!-- grade_id: cell-7881cdc32535f82e, points: 0, grade: false, solution: true -->
```python
growth_rate_pwt.sort_values()
growth_rate_pwt.tail(10)
```

Output:
```
2322   NaN
2324   NaN
2338   NaN
2350   NaN
2362   NaN
2370   NaN
2372   NaN
2386   NaN
2400   NaN
2414   NaN
Name: gdp_pc, dtype: float64
```

Now in WDI

<!-- grade_id: cell-5413cde488f968b7, points: 0, grade: false, solution: true -->
```python
growth_rate_wdi.sort_values()
growth_rate_wdi.tail(10)
```

Output:
```
259         NaN
260    1.627125
261   -0.010236
262    0.565360
263    0.414442
264         NaN
265         NaN
266    0.153261
267    0.409522
268   -0.175563
dtype: float64
```

# 7. Convergence

Generate a `plotly` scatter plot that shows the income level in the earliest period on the horizontal axis and the growth rate between earliest and latest in the vertical axis. Make sure hovering on a symbols shows the information of the country. Use colors for WB regions and shapes for WB income groups. Export it to file `dataset_convergence_early_late.pdf`. Do you see anything noteworthy?

Start by creating a dataframe with only data for the earliest date `dataset_early` and then merge it with the growth rates dataframe `growth_dataset` (make sure to add a suffix to distinguish variables if needed). Drop any NA's to simplify plotting.

First PWT

<!-- grade_id: cell-4ea1badb9248e4f2, points: 0, grade: false, solution: true -->
```python
#This is as far as I could get..

```

Output:
```
NameError                                 Traceback (most recent call last)
Input In [429], in <cell line: 1>()
----> 1 growth_rates_pwt.plot()

NameError: name 'growth_rates_pwt' is not defined
```

Now WDI

```python
wdi_early = wdi.loc[wdi.year==year_min_wdi].reset_index(drop=True).copy()
wdi_early = wdi_early.merge(growth_wdi, on=['country_code', 'country_name'], suffixes=['', '_growth'])
wdi_early = wdi_early.merge(wbcountries[['iso3c', 'region', 'incomeLevel']], left_on='country_code', right_on='iso3c')
wdi_early = wdi_early.dropna().reset_index(drop=True)
wdi_early
```

Now let's plot this data for PWT

<!-- grade_id: cell-00b095ab7125e6b5, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
Input In [428], in <cell line: 2>()
      1 # YOUR CODE HERE
----> 2 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-0e6126c984bab237, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'pwt_convergence_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf'))
```

Now for WDI

<!-- grade_id: cell-c91a2028c33611e6, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-cbf1b575e9cddb9d, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'wdi_convergence_' + str(year_min_wdi) + '_' + str(year_max_wdi) + '.pdf'))
```

# 8. Persistence

Generate a `plotly` scatter plot that shows the income level in the earliest period on the horizontal axis and the income level in the latest in the vertical axis. Make sure hovering on a symbols shows the information of the country. Use colors for WB regions and shapes for WB income groups. Export it to file `dataset_persistence_early_late.pdf`. Do you see anything noteworthy?

Start by creating a dataframe with only data for the latest date `dataset_late` and then merge it with the `dataset_early` dataframe you cerated in teh previous point (make sure to add a suffix to distinguish variables if needed). Drop any NA's to simplify plotting.

Start with PWT

<!-- grade_id: cell-187179cae74cdcbd, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now WDI

<!-- grade_id: cell-f88aac24a08de9b4, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

Now let's plot this data for PWT

<!-- grade_id: cell-90b1bf6efdd50558, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-3434c9dc87ec5a10, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'pwt_persistence_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf'))
```

Now WDI

<!-- grade_id: cell-0f954ffbfb90b55a, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-7ffcfe5a3ccd96ef, points: 5, grade: true, solution: false -->
```python
assert(os.path.exists(pathgraphs + 'wdi_persistence_' + str(year_min_pwt) + '_' + str(year_max_pwt) + '.pdf'))
```